# Exercise 2 - Lists and indices

The exercise for this week is meant to help you better understand data types and lists in Python, and practice saving changes to your files using Git and GitHub. 

Exercise 2 consists of 3 problems that you can find in this notebook. For each problem you need to modify the notebook either by fixing problems in existing code, or by adding your own solutions. Remember to push your changes to GitHub already while working, not only at the very end!

- **Exercise 2 is due by 09:00 on 18.9.**
- Scores on this exercise are out of **20 points**.

## Please note

- **Some of the cells in this Jupyter Notebook are "Read-only" which means that you are not able to edit them.**

- **Some of the code cells contain tests. If these tests fail, you know there is still something wrong with your code.**

- **Some of the code cells raise** `NotImplementedError`. **This error tells you that you need to somehow modify that code cell.**


## Where to find help

If you're having trouble, please remember you can

- Come to the work sessions on Thursday and/or Friday for help from the course assistants
- Check out [the hints for this week's exercise](https://geo-python.github.io/site/lessons/L2/exercise-2.html) on the course webpage
- Post questions to Slack

## Problem 0 - Cloning Exercise repository to JupyterLab (*0 points*)

Before starting to work with the actual problems for this week, you should start a new JupyterLab instance and clone **your own** Exercise 2 repository (e.g. `exercise-2-htenkanen`) into the instance using **Git** as we saw in the [lecture this week](https://geo-python.github.io/site/lessons/L2/git-basics.html). 

## Problem 1 - Making changes to code and using Git (*7 points*)

Your first task for this week is to make some modifications to the Python code in the cells below to fix the errors it gives once you run the code cells.
The code should allow users to find the operational time of an [FMI observation station](http://en.ilmatieteenlaitos.fi/observation-stations) by setting the `selected_station` variable.
*Don't worry about the case of a station name that is not on the list.*

**The main aim is to print out an informative sentence that tells how long the selected station has been operational, for example** 

```
'The Helsinki Harmaja station has been operational for 30 years.'
```
or
```
'The Helsinki Kaivopuisto station has been operational for 115 years.'
```

However, if you try to run all the code cells under Problem 1, you will run into errors. Your task is to find **3 things to fix** in the code cells under Problem 1. Once you succeed, the final print statement should work, and you should be able to modify the output only by setting a different station name.


### Grading for Problem 1

Your score for this problem will be based on

- Fixing **3 problems in the code** to get the code working as expected
- Committing each change separately to GitHub using Git
- Listing the changes you needed to make to get the code working in plain English by modifying the existing text or adding a Markdown cell above each change

Here we create and fill lists of station names and starting years for observation.
**NOTE**: Long lines can be split into multiple lines using the "\\" character.

*These code cells are "Read-only" - you won't be able to modify the input lists, but you still need to run them.*

```python
# Read-only code cell
station_names = ['Harmaja', 'Kaisaniemi', 'Kaivopuisto', 'Kumpula', 'lighthouse', \
                'Malmi airfield', 'Suomenlinna aaltopoiju', 'Vuosaari harbour']
```

```python
# Read-only code cell
station_start_years = [1989, 1844, 1904, 2005, 2003, 1937, 2016, 2012]
```

Below we can define the selected station:

<!-- grade_id: selected_station, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

selected_station = None
```

<!-- grade_id: selected_station_test, points: 1, grade: true, solution: false -->
```python
print(selected_station)
```

Next we find the location of the selected station in the list of station names:

```python
# Read-only code cell
station_index = station_names.index(selected_station)
```

Now we can calculate the years the station has been operational:

<!-- grade_id: station_years, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

station_years = 2019 - station_start_years[selected_station]
```


Now, we can print the station name and number of years of operation to the screen.
A sample print statement on screen is `'The Helsinki Harmaja station has been operational for 30 years.'` The print statements should work for all listed stations only by modifying the `selected_station` -variable that is defined earlier in this problem.

<!-- grade_id: print_statement, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

print_statement = "The Helsinki " + selected_station + " station has been operational for "+ \
      str(station_years)+ " years.

print(print_statement)
```

#### Tests for problem 1

Finally, we can run some tests to see if things are working as they should! These tests may help you in solving the problem. Note, we also run some hidden tests for checking the exercises. If these tests pass, it's not a quarantee of a 100 % correct answer!

<!-- grade_id: station_years_test, points: 1, grade: true, solution: false -->
```python
# importing tools for running some tests
from nose.tools import ok_, assert_equal

#Validate that station_years is a variable of integer type
assert_equal(type(station_years),int)
```

<!-- grade_id: print_statement_test1, points: 2, grade: true, solution: false -->
```python
#Validate that the print statement is correct for Harmaja station:
# Note! Your print statement should work for all the stations!
assert_equal(print_statement, 'The Helsinki Harmaja station has been operational for 30 years.')

```

#### Summarize what you did in problem 1 using markdown

Describe the 3 code changes you made below in plain English (you can, for example, use bullet points here):

<!-- grade_id: final_answer_question_test, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

## Problem 2 - Accessing monthly average temperatures (*9 points*)

The table below presents [monthly average temperatures recorded at the Helsinki Malmi airport](https://www.timeanddate.com/weather/finland/helsinki/climate).

| Month     | Temperature [°C] |
| --------- | :--------------: |
| January   | -3.5             |
| February  | -4.5             |
| March     | -1.0             |
| April     | 4.0              |
| May       | 10.0             |
| June      | 15.0             |
| July      | 18.0             |
| August    | 16.0             |
| September | 11.5             |
| October   | 6.0              |
| November  | 2.0              |
| December  | -1.5             |

In the code cells below, write some Python code that allows users to select a month and have the monthly average temperature printed to the screen.
For example, your code should display the following for the month of March:

```
The average temperature in Helsinki in March is -1.0
```

### Grading for Problem 2

Your score for this problem will be based on

- Having your notebook display the monthly average temperature in a selected month, set by defining the variable `selected_month_index`.
    - **NOTE**: We expect you to use the index value to select the month, not the name of the month (e.g., "May").
- Having it work for all 12 months in the year.
- Describe how your code works in a few sentences of plain English in the Markdown cell above your code
- Including inline comments to clearly explain how the code works between most and/or all lines of your code in its Python cell
- Pushing your script to your GitHub repository for this week's lesson.

<!-- grade_id: define_lists, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

# here is the variable which you use to set the selected month:
selected_month_index = None

# Here are also the names of the two lists you should use for this problem. Please don't change the variable names!
months = None
average_temp = None

# Using the lists and their indices, generate the desired print statement:
print_statement = None
```

<!-- grade_id: define_lists_test, points: 2, grade: true, solution: false -->
```python
# Check your print statement:
print(print_statement)
```

#### Tests for problem 2

Let's run some additional tests to see if things are working as they should! These tests may help you in solving the problem. Note, we also run some hidden tests for checking the exercises. If these tests pass, it's not a quarantee of a 100 % correct answer!

<!-- grade_id: list_length_test, points: 2, grade: true, solution: false -->
```python
#Validate the length of two lists are 12
ok_(len(months)==12)
ok_(len(average_temp)==12)
```

<!-- grade_id: list_type_test, points: 2, grade: true, solution: false -->
```python
#Validate that variable months and average_temp are lists
ok_(isinstance(months, list), 'Variable months is not a list')
ok_(isinstance(average_temp, list), 'Variable average_temp is not a list')
```

<!-- grade_id: print_statement_test, points: 3, grade: true, solution: false -->
```python
#Validate the print statement is correct; 
# Set selected_month_index to correspond with July before running this cell.
# Note! Your code should work with any of the 12 months!
assert_equal(print_statement, 'The average temperature in Helsinki in July is 18.0')

```

## Problem 3 - More practice using Markdown (*4 points*)

The last task in this week's exercise is to answer some questions and add an image in this notebook using Markdown. **Note**: You may want to read a bit more about [formatting text in Github-flavored Markdown](https://help.github.com/articles/basic-writing-and-formatting-syntax/).


### Grading for Problem 3

Your score for this problem will be based on

- Your answers to the three questions that are part of this problem
- Posting an image of a favorite animal using Markdown


#### 1. Give your responses to these three questions about this week's lesson:

    - What did you learn?
    - What was unclear?
    - What would you change?
    
    *Please use Markdown lists when answering these questions.*
        

<!-- grade_id: answer_questions, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

#### 2. Adding an image

Just for fun, add an image of an animal that you like along with a short caption giving its name and anything special you might like to add.

You should add an image by linking to a website with the image, or if by uploading an image to your GitHub repository and linking to that (if you are looking for a challenge).

We suggest that you search for images in a repository that includes licencing information such as [Wikimedia Commons](https://commons.wikimedia.org/wiki/Main_Page) or [Pixabay](https://pixabay.com/). You are, of course, also welcome to upload your own animal images. You can add it in a new Markdown cell under those used to answer the questions in part 1 of this problem.

<!-- grade_id: insert_image, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE